# Basic linear algebra in Julia
Author: Andreas Noack Jensen (MIT & JuliaComputing) (https://twitter.com/anoackjensen?lang=en)
(with edits from Jane Herriman)

First let's define a random matrix

```julia
#
# Return a 3x3 matrix with random numbers in the range 1..4
# 
A = rand(1:4,3,3)
```

Output:
```
3×3 Matrix{Int64}:
 1  2  2
 4  4  1
 4  4  4
```

# References

## Norm of a matrix

[https://docs.google.com/presentation/d/1nrcj_vZ52r16klTf3uOLr_xF52WPxg3vx7r28KKguZI/edit#slide=id.g2d41f2b6d9d_0_0]

## LU Factorization

[https://docs.google.com/presentation/d/1nrcj_vZ52r16klTf3uOLr_xF52WPxg3vx7r28KKguZI/edit#slide=id.g2d41f2b6d9d_0_12]

## QR Factorization

[https://docs.google.com/presentation/d/1nrcj_vZ52r16klTf3uOLr_xF52WPxg3vx7r28KKguZI/edit#slide=id.g2d41f2b6d9d_0_19]


## Orthogonal Matrix

[https://docs.google.com/presentation/d/1nrcj_vZ52r16klTf3uOLr_xF52WPxg3vx7r28KKguZI/edit#slide=id.g2d42a15a002_0_10]


Define a vector of size 3, filled with ones

```julia
x = fill(1.0, (3,)) # = fill(1.0, 3)
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

Notice that $A$ has type Array{Int64,2} but $x$ has type Array{Float64,1}. Julia defines the aliases Vector{Type}=Array{Type,1} and Matrix{Type}=Array{Type,2}. 

Many of the basic operations are the same as in other languages
#### Multiplication

```julia
b = A*x
```

Output:
```
3-element Vector{Float64}:
  5.0
  9.0
 12.0
```

#### Transposition
As in other languages `A'` is the conjugate transpose, or adjoint

```julia
A'
```

Output:
```
3×3 adjoint(::Matrix{Int64}) with eltype Int64:
 1  4  4
 2  4  4
 2  1  4
```

and we can get the transpose with

```julia
transpose(A)
```

Output:
```
3×3 transpose(::Matrix{Int64}) with eltype Int64:
 1  4  4
 2  4  4
 2  1  4
```

#### Transposed multiplication
Julia allows us to write this without *

```julia
A'A
```

Output:
```
3×3 Matrix{Int64}:
 33  34  22
 34  36  24
 22  24  21
```

#### Solving linear systems 
The problem $Ax=b$ for ***square*** $A$ is solved by the  function
```
\
```


```julia
A\b
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

`A\b` gives us the *least squares solution* if we have an overdetermined linear system (a "tall" matrix)

```julia
Atall = rand(3, 2)
```

Output:
```
3×2 Matrix{Float64}:
 0.902659  0.808179
 0.501668  0.606486
 0.798685  0.513372
```

```julia
Atall\b
```

Output:
```
2-element Vector{Float64}:
 12.397425007219077
 -1.7434522679536013
```

and the *minimum norm least squares solution* if we have a rank-deficient least squares problem

```julia
v = rand(3)
rankdef = hcat(v, v)
```

Output:
```
3×2 Matrix{Float64}:
 0.750717  0.750717
 0.843297  0.843297
 0.826987  0.826987
```

```julia
rankdef\b
```

Output:
```
2-element Vector{Float64}:
 5.429067111418461
 5.429067111418458
```

Julia also gives us the minimum norm solution when we have an underdetermined solution (a "short" matrix)

```julia
bshort = rand(2)
Ashort = rand(2, 3)
```

Output:
```
2×3 Matrix{Float64}:
 0.273918  0.774197  0.251048
 0.197762  0.982008  0.759963
```

```julia
Ashort\bshort
```

Output:
```
3-element Vector{Float64}:
 -4.139674407706176
 -0.9797338415813472
 23.563074056363842
```

# The LinearAlgebra library

While much of linear algebra is available in Julia by default (as shown above), there's a standard library named `LinearAlgebra` that brings in many more relevant names and functions. In particular, it provides factorizations and some structured matrix types.  As with all packages, you can bring these additional features into your session with a `using LinearAlgebra`.

### Exercises

#### 10.1 
Take the inner product (or "dot" product) of a vector `v` with itself and assign it to variable `dot_v`.



```julia
v = [1,2,3]
```

Output:
```
3-element Vector{Int64}:
 1
 2
 3
```

```julia

```

<!-- grade_id: cell-913fef9b0d19cd52, points: 1, grade: true, solution: false -->
```julia
@assert dot_v == 14
```

#### 10.2 
Take the outer product of a vector v with itself and assign it to variable `outer_v`

```julia

```

```julia
@assert outer_v == [1 2 3
                    2 4 6
                    3 6 9]
```

#### 10.3 
Use [LinearAlgebra.cross](https://docs.julialang.org/en/v1/stdlib/LinearAlgebra/#LinearAlgebra.cross) to compute the cross product of a vector v with itself and assign it to variable `cross_v`

```julia

```

<!-- grade_id: cell-e6b6970ffe104df5, points: 1, grade: true, solution: false -->
```julia
@assert cross_v == [0, 0, 0]
```